// Numbas version: finer_feedback_settings
{"name": "Tallregning", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": true, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Produkt", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Produkt 1", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om produkt.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "b1", "b2", "c1", "c2", "c3", "d1", "d2", "d3"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": false, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{a1}\\cdot \\var{a2}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: $- \\cdot + = -$</p>"}], "minValue": "a1*a2", "maxValue": "a1*a2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{b1}\\cdot(\\var{b2})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: $- \\cdot - = +$</p>"}], "minValue": "b1*b2", "maxValue": "b1*b2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Produkt 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om produkt.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "b1", "b2", "c1", "c2", "c3", "d1", "d2", "d3"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{c1}\\cdot \\var{c2} \\cdot(\\var{c3})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: $-\\cdot +\\cdot -=(-\\cdot +)\\cdot -=-\\cdot -=+$</p>"}], "minValue": "c1*c2*c3", "maxValue": "c1*c2*c3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{d1}\\cdot (\\var{d2}) \\cdot(\\var{d3})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: $-\\cdot-\\cdot- = (-\\cdot-)\\cdot- = +\\cdot-=-$</p>"}], "minValue": "d1*d2*d3", "maxValue": "d1*d2*d3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Sum og produkt", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Sum + produkt niv\u00e5 1", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om sum og produkt.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c4": {"name": "c4", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d4": {"name": "d4", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e4": {"name": "e4", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a3", "a2", "b1", "b2", "b3", "c1", "c2", "c3", "c4", "d1", "d2", "d3", "d4", "e1", "e2", "e3", "e4"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{a1} + \\var{a2}\\cdot \\var{a3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>&nbsp;parentesuttrykk</li>\n<li>&nbsp;potens</li>\n<li>&nbsp;<strong>multiplikasjon</strong> og divisjon</li>\n<li>&nbsp;<strong>addisjon</strong> og subtraksjon&nbsp;</li>\n</ol>\n<p>F&oslash;rst multiplikasjon, deretter addisjon.</p>"}], "minValue": "a1+a2*a3", "maxValue": "a1+a2*a3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{b1} + \\var{b2}\\cdot (\\var{b3})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ul>\n<li>F&oslash;rst multiplikasjon,</li>\n<li>deretter addisjon.</li>\n</ul>\n<p>Husk: $+ \\cdot - = -$</p>\n<ol></ol>\n<ul></ul>"}], "minValue": "b1+b2*b3", "maxValue": "b1+b2*b3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Sum + produkt niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om sum og produkt.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c4": {"name": "c4", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d4": {"name": "d4", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e4": {"name": "e4", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a3", "a2", "b1", "b2", "b3", "c1", "c2", "c3", "c4", "d1", "d2", "d3", "d4", "e1", "e2", "e3", "e4"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{c1}\\cdot \\var{c2} + \\var{c3}\\cdot \\var{c4}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ul>\n<li>Multipliser begge produktene f&oslash;rst,&nbsp;</li>\n<li>deretter addisjon.</li>\n</ul>"}], "minValue": "c1*c2+c3*c4", "maxValue": "c1*c2+c3*c4", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{d1}\\cdot \\var{d2} + (\\var{d3})\\cdot \\var{d4}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ul>\n<li>Multipliser begge produktene f&oslash;rst,&nbsp;</li>\n<li>deretter addisjon.</li>\n</ul>\n<p>Husk: $+ \\cdot - = -$</p>"}], "minValue": "d1*d2+d3*d4", "maxValue": "d1*d2+d3*d4", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{e1}\\cdot \\var{e2} - \\var{e3}\\cdot (\\var{e4})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ul>\n<li>Multipliser begge produktene f&oslash;rst,&nbsp;</li>\n<li>deretter addisjon.</li>\n</ul>\n<p>Husk fortegnsregler:</p>\n<ul>\n<li>$+ \\cdot - = -$</li>\n<li>$- \\cdot - = +$</li>\n</ul>"}], "minValue": "e1*e2-e3*e4", "maxValue": "e1*e2-e3*e4", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Divisjon", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Divisjon", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om divisjon.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "a2*ak", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "ak": {"name": "ak", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "-b2*bk", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "bk": {"name": "bk", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "c3list[j]", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "c2list[j]", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "c1list[j]", "description": "", "templateType": "anything", "can_override": false}, "c1list": {"name": "c1list", "group": "Ungrouped variables", "definition": "[2,3,4,5,6]", "description": "", "templateType": "anything", "can_override": false}, "j": {"name": "j", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "c2list": {"name": "c2list", "group": "Ungrouped variables", "definition": "[3,5,6,4,2]", "description": "", "templateType": "anything", "can_override": false}, "c3list": {"name": "c3list", "group": "Ungrouped variables", "definition": "[6,5,4,4,2]", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[4,5,6,10,7]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[6, 8, 5, 3, 9]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[3,4,10, 6, 3]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "ak", "b1", "b2", "bk", "c3", "c2", "c1", "c1list", "j", "c2list", "c3list", "d1", "d2", "d3", "d1list", "d2list", "d3list", "dk"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{a1} : \\var{a2}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\var{a1}: \\var{a2}$ betyr $\\var{a1}$ delt p&aring; $\\var{a2}$.&nbsp;</p>\n<p>Annen skrivem&aring;te er $\\frac{\\var{a1}}{\\var{a2}}$</p>"}], "minValue": "ak", "maxValue": "ak", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{b1} : \\var{b2}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Fortegnsregler gjelder b&aring;de multiplikasjon og divisjon. $- : + = -$</p>"}], "minValue": "-bk", "maxValue": "-bk", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Divisjon og multiplikasjon", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om divisjon.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "a2*ak", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "ak": {"name": "ak", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "-b2*bk", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "bk": {"name": "bk", "group": "Ungrouped variables", "definition": "random(2..4)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "c3list[j]", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "c2list[j]", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "c1list[j]", "description": "", "templateType": "anything", "can_override": false}, "c1list": {"name": "c1list", "group": "Ungrouped variables", "definition": "[2,3,4,5,6]", "description": "", "templateType": "anything", "can_override": false}, "j": {"name": "j", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "c2list": {"name": "c2list", "group": "Ungrouped variables", "definition": "[3,5,6,4,2]", "description": "", "templateType": "anything", "can_override": false}, "c3list": {"name": "c3list", "group": "Ungrouped variables", "definition": "[6,5,4,4,2]", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[4,5,6,10,7]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[6, 8, 5, 3, 9]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[3,4,10, 6, 3]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "ak", "b1", "b2", "bk", "c3", "c2", "c1", "c1list", "j", "c2list", "c3list", "d1", "d2", "d3", "d1list", "d2list", "d3list", "dk"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{c1}\\cdot \\var{c2} : \\var{c3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>Parentesuttrykk</li>\n<li>Potens</li>\n<li><strong>Multiplikasjon</strong> og <strong>divisjon</strong></li>\n<li>Addisjon og subtraksjon</li>\n</ol>\n<p>Multiplikasjon og divisjon beregnes samtidig. Multipliser f&oslash;rst, s&aring; del etterp&aring;. Kan ogs&aring; gj&oslash;re omvendt.&nbsp;</p>"}], "minValue": "c1*c2/c3", "maxValue": "c1*c2/c3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{d1}\\cdot \\var{d2} : \\var{d3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Multiplikasjon og divisjon beregnes samtidig. Multipliser f&oslash;rst, s&aring; del etterp&aring;. Kan ogs&aring; gj&oslash;re omvendt.&nbsp;</p>"}], "minValue": "d1*d2/d3", "maxValue": "d1*d2/d3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Parentes", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", ""], "variable_overrides": [[], []], "questions": [{"name": "Parentes og multiplikasjon", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om parentes.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[28, 30, 40, 27, 50]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[9, 4, 10, 5, 8]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[-2, 2, -6, 4, -3]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "e1list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "e2list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "e3list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Ungrouped variables", "definition": "[14, 11, 25, 14, 28]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Ungrouped variables", "definition": "[21, 13, -7, 16, -4]", "description": "", "templateType": "anything", "can_override": false}, "e3list": {"name": "e3list", "group": "Ungrouped variables", "definition": "[7, 6, 3, 5, 4]", "description": "", "templateType": "anything", "can_override": false}, "ek": {"name": "ek", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "a3", "b1", "b2", "b3", "c1", "c2", "c3", "d1", "d3", "d1list", "d2list", "d2", "d3list", "dk", "e1", "e2", "e3", "e1list", "e2list", "e3list", "ek"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(\\var{a1}+ \\var{a2}) \\cdot \\var{a3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li><strong>Parentesuttrykk</strong></li>\n<li>Potens</li>\n<li><strong>Multiplikasjon</strong> og divisjon</li>\n<li>Addisjon og subtraksjon</li>\n</ol>\n<p>L&oslash;s f&oslash;rst uttrykket inni parentesen. Deretter multipliser med $\\var{a3}$</p>"}], "minValue": "(a1+a2)*a3", "maxValue": "(a1+a2)*a3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{b1}\\cdot (\\var{b2} + \\var{b3})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Parentesuttrykk f&oslash;r multiplikasjon</p>"}], "minValue": "b1*(b2+b3)", "maxValue": "b1*(b2+b3)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{c1}\\cdot (\\var{c2} + (\\var{c3}))$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Parentesuttrykk f&oslash;r multiplikasjon.</p>\n<p>Husk fortegnsregel: $+ \\cdot - = -$</p>"}], "minValue": "c1*(c2+c3)", "maxValue": "c1*(c2+c3)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Parentes og divisjon", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om parentes.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[28, 30, 40, 27, 50]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[9, 4, 10, 5, 8]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[-2, 2, -6, 4, -3]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "e1list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "e2list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "e3list[ek]", "description": "", "templateType": "anything", "can_override": false}, "e1list": {"name": "e1list", "group": "Ungrouped variables", "definition": "[14, 11, 25, 14, 28]", "description": "", "templateType": "anything", "can_override": false}, "e2list": {"name": "e2list", "group": "Ungrouped variables", "definition": "[21, 13, -7, 16, -4]", "description": "", "templateType": "anything", "can_override": false}, "e3list": {"name": "e3list", "group": "Ungrouped variables", "definition": "[7, 6, 3, 5, 4]", "description": "", "templateType": "anything", "can_override": false}, "ek": {"name": "ek", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "a3", "b1", "b2", "b3", "c1", "c2", "c3", "d1", "d3", "d1list", "d2list", "d2", "d3list", "dk", "e1", "e2", "e3", "e1list", "e2list", "e3list", "ek"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$(\\var{e2}+\\var{e1}): \\var{e3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Parentesene f&oslash;r divisjon</p>"}], "minValue": "(e1+e2)/e3", "maxValue": "(e1+e2)/e3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$\\var{d1}: (\\var{d3}+\\var{d2})$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Parentesene f&oslash;r divisjon</p>"}], "minValue": "d1/(d2+d3)", "maxValue": "d1/(d2+d3)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}, {"name": "Oppsummering", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", ""], "variable_overrides": [[], [], []], "questions": [{"name": "Oppsummering niv\u00e5 1", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om regnerekkef&oslash;lge.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6 except -a1)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(2..6 except b2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-4..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[-4,6,-12,-2,3]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[10,2,2,11,4]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[3,4,5,3,7]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "random(6..10)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e4": {"name": "e4", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e5": {"name": "e5", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "a3", "b2", "b1", "b3", "c1", "c2", "c3", "d1", "d2", "d3", "d1list", "d2list", "d3list", "dk", "e1", "e2", "e3", "e4", "e5"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $(\\var{a1}+\\var{a2})^2+ \\var{a3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>&nbsp;<strong>parentesuttrykk</strong></li>\n<li>&nbsp;<strong>potens</strong></li>\n<li>&nbsp;multiplikasjon og divisjon</li>\n<li>&nbsp;<strong>addisjon</strong> og subtraksjon&nbsp;</li>\n</ol>\n<p>L&oslash;s f&oslash;rst uttrykket inni parentesen. Deretter opph&oslash;yd resultatet i andre; til slutten, legg til $\\var{a3}$</p>"}], "minValue": "(a1+a2)^2+a3", "maxValue": "(a1+a2)^2+a3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{b1}+ (\\var{b2}-\\var{b3})^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk: Parentesuttrykk f&oslash;r potens; potens f&oslash;r subtraksjon.</p>"}], "minValue": "b1 + (b2-b3)^2", "maxValue": "b1 + (b2-b3)^2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Oppsummering niv\u00e5 2", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om regnerekkef&oslash;lge.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6 except -a1)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(2..6 except b2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-4..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[-4,6,-12,-2,3]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[10,2,2,11,4]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[3,4,5,3,7]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "random(6..10)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e4": {"name": "e4", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e5": {"name": "e5", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "a3", "b2", "b1", "b3", "c1", "c2", "c3", "d1", "d2", "d3", "d1list", "d2list", "d3list", "dk", "e1", "e2", "e3", "e4", "e5"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $\\var{c1}\\cdot \\var{c2} + (\\var{c3})^3$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>Parentesuttrykk</li>\n<li><strong>Potens</strong></li>\n<li><strong>Multiplikasjon</strong> og divisjon</li>\n<li><strong>Addisjon</strong> og subtraksjon</li>\n</ol>\n<p>Beregn f&oslash;rst potensen. Deretter multipliser $\\var{c1}$ med $\\var{c2}.$ Til slutten, beregn summen.</p>\n<p></p>"}], "minValue": "c1*c2 + c3^3", "maxValue": "c1*c2 + c3^3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$(\\var{d1}+\\var{d2})^2:\\var{d3}$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li><strong>Parentesuttrykk</strong></li>\n<li><strong>Potens</strong></li>\n<li>Multiplikasjon og <strong>divisjon</strong></li>\n<li>Addisjon og subtraksjon</li>\n</ol>\n<p>Beregn f&oslash;rst parentesuttrykk. Deretter beregn potensen. Til slutten, del p&aring; $\\var{d3}$.</p>"}], "minValue": "(d1+d2)^2/d3", "maxValue": "(d1+d2)^2/d3", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}, {"name": "Oppsummering niv\u00e5 3", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<h3>L&oslash;s de f&oslash;lgende oppgavene om regnerekkef&oslash;lge.</h3>", "advice": "", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(-6..-2)", "description": "", "templateType": "anything", "can_override": false}, "a2": {"name": "a2", "group": "Ungrouped variables", "definition": "random(2..6 except -a1)", "description": "", "templateType": "anything", "can_override": false}, "a3": {"name": "a3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b2": {"name": "b2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "b3": {"name": "b3", "group": "Ungrouped variables", "definition": "random(2..6 except b2)", "description": "", "templateType": "anything", "can_override": false}, "c1": {"name": "c1", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c2": {"name": "c2", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "c3": {"name": "c3", "group": "Ungrouped variables", "definition": "random(-4..-2)", "description": "", "templateType": "anything", "can_override": false}, "d1": {"name": "d1", "group": "Ungrouped variables", "definition": "d1list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d2": {"name": "d2", "group": "Ungrouped variables", "definition": "d2list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d3": {"name": "d3", "group": "Ungrouped variables", "definition": "d3list[dk]", "description": "", "templateType": "anything", "can_override": false}, "d1list": {"name": "d1list", "group": "Ungrouped variables", "definition": "[-4,6,-12,-2,3]", "description": "", "templateType": "anything", "can_override": false}, "d2list": {"name": "d2list", "group": "Ungrouped variables", "definition": "[10,2,2,11,4]", "description": "", "templateType": "anything", "can_override": false}, "d3list": {"name": "d3list", "group": "Ungrouped variables", "definition": "[3,4,5,3,7]", "description": "", "templateType": "anything", "can_override": false}, "dk": {"name": "dk", "group": "Ungrouped variables", "definition": "random(0..4)", "description": "", "templateType": "anything", "can_override": false}, "e1": {"name": "e1", "group": "Ungrouped variables", "definition": "random(6..10)", "description": "", "templateType": "anything", "can_override": false}, "e2": {"name": "e2", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e3": {"name": "e3", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "e4": {"name": "e4", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "e5": {"name": "e5", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "f1": {"name": "f1", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "f2": {"name": "f2", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "f3": {"name": "f3", "group": "Ungrouped variables", "definition": "random(2..5)", "description": "", "templateType": "anything", "can_override": false}, "f4": {"name": "f4", "group": "Ungrouped variables", "definition": "random(2..6)", "description": "", "templateType": "anything", "can_override": false}, "f5": {"name": "f5", "group": "Ungrouped variables", "definition": "random(5..10)", "description": "", "templateType": "anything", "can_override": false}, "f6": {"name": "f6", "group": "Ungrouped variables", "definition": "random(4..7)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["a1", "a2", "a3", "b2", "b1", "b3", "c1", "c2", "c3", "d1", "d2", "d3", "d1list", "d2list", "d3list", "dk", "e1", "e2", "e3", "e4", "e5", "f1", "f2", "f3", "f4", "f5", "f6"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $((\\var{e1}-\\var{e2})\\cdot \\var{e3} - \\var{e4}\\cdot \\var{e5})^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>Parentesuttrykk</li>\n<li>Potens</li>\n<li>Multiplikasjon og divisjon</li>\n<li>Addisjon og subtraksjon</li>\n</ol>\n<p>Pass p&aring; at det er to parentesuttrykk! Start med den indre parantesen.</p>\n<p>N&aring;r det gjelder uttrykket i den ytre parentesen (dvs. $(\\var{e1}-\\var{e2})\\cdot \\var{e3} - \\var{e4}\\cdot \\var{e5}$) m&aring; vi f&oslash;lge regnerekkef&oslash;lge igjen. Vi opph&oslash;yder i andre helt til slutten!</p>"}], "minValue": "((e1-e2)*e3-e4*e5)^2", "maxValue": "((e1-e2)*e3-e4*e5)^2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Beregn $((\\var{f1}+\\var{f2})^2+ \\var{f3})\\cdot \\var{f4} - \\var{f5}\\cdot \\var{f6}^2$</p>", "stepsPenalty": 0, "steps": [{"type": "information", "useCustomName": true, "customName": "Tips", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Husk regnerekkef&oslash;lge:</p>\n<ol>\n<li>Parentesuttrykk</li>\n<li>Potens</li>\n<li>Multiplikasjon og divisjon</li>\n<li>Addisjon og subtraksjon</li>\n</ol>\n<p>Pass p&aring; at det er to parentesuttrykk! Start med den indre parantesen.</p>\n<p>Pass ogs&aring; p&aring; at $\\var{f5} \\cdot \\var{f6}^2 \\neq (\\var{f5} \\cdot \\var{f6})^2$</p>"}], "minValue": "((f1+f2)^2+f3)*f4-f5*f6^2", "maxValue": "((f1+f2)^2+f3)*f4-f5*f6^2", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "showAllParts": false}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "typeendtoleave": false, "startpassword": "", "autoSubmit": true, "allowAttemptDownload": false, "downloadEncryptionKey": "", "showresultspage": "oncompletion"}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"enterreviewmodeimmediately": true, "showactualmarkwhen": "always", "showtotalmarkwhen": "always", "showanswerstatewhen": "always", "showpartfeedbackmessageswhen": "always", "showexpectedanswerswhen": "inreview", "showadvicewhen": "inreview", "allowrevealanswer": true, "intro": "", "end_message": "", "results_options": {"printquestions": true, "printadvice": true}, "feedbackmessages": [], "reviewshowexpectedanswer": true, "showanswerstate": true, "reviewshowfeedback": true, "showactualmark": true, "showtotalmark": true, "reviewshowscore": true, "reviewshowadvice": true}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "extensions": [], "custom_part_types": [], "resources": []}